```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

# Homework 02: Arrays and Tables

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to load the provided modules. Each time you start your server, you will need to execute this cell again to load the modules.

```python
# Don't change this cell; just run it. 
import numpy as np
from datascience import *
```

**Helpful Resource:**
- [Python Reference](http://data8.org/sp22/python-reference.html): Cheat sheet of helpful array & table methods used in this course.

**Recommended Reading:**
- [What is Data Science](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causality and Experiments](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programming in Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

**Note: This homework has hidden tests on it. Additional tests will be run once your homework is submitted for grading. While you may pass all the tests you have access to before submission, you may not earn full credit if you do not pass the hidden tests as well.**. 

Many of the tests you have access to before submitting only test to ensure you have given an answer that is formatted correctly and/or you have given an answer that *could* make sense in context. For example, a test you have access to while completing the assignment may check that you selected a valid choice for a multiple choice problem (1, 2, or 3) or that your answer is an integer between 0 and 50 if asked to count a subset of states in the United States. The tests that are run after submission will evaluate your work for accuracy. **Do not assume that just because all your tests pass before submission means that your answers are correct!**

Consult with your teacher and course syllabus for information and policies regarding appropriate collaboration with other students, appropriate use of AI tools, and submission of late work.

## 1. Creating Arrays

### Question 1.1

Make an array called `rainbow` containing the following seven strings: "Red", "Orange", "Yellow", "Green", "Blue", "Indigo", "and Violet".

```python
rainbow = ...
rainbow
```

```python
grader.check("q1_1")
```

### Question 1.2
2

Make an **array** called `just_numbers` containing the following numbers (in the given order):

1. $2^{\tan(13)}$
2. $16^{0.25}$
3. $\cos(12)$
4. $-8$
   

**Hint:** `tan` and `cos` are functions in the `math` module. Importing modules is covered in a previous lab! You should import the module in your solution cell below.

**Note:** Python lists are different/behave differently than numpy arrays. In this course, we use `numpy` arrays, so please make an **array**, not a python list.

```python
# Our solution involved one extra line of code before creating
# just_numbers.
...
just_numbers = ...
just_numbers
```

```python
grader.check("q1_2")
```

### `join`
All strings in Python have a built-in method called `join`.  The `join` method takes only one argument as an input, a collection (like an array) of strings. The `join` method returns a single string as an output. 

The output string of `a_string.join(an_array)` is a single string that's the [concatenation](https://en.wikipedia.org/wiki/Concatenation) ("putting together") of all the strings in `an_array`, with `a_string` inserted in between each string from the array. Run the cell below to see it in action:

```python
area_code = '828'
prefix = '347'
line_number = '9100'

phone_number = '-'.join(make_array(area_code, prefix, line_number))
phone_number
```

Notice how the single dash, `-`, was inserted between each of the three strings contained in the array to generate the phone number for the NCSSM-Morganton front desk.

### Question 1.3

Use the array `rainbow` and the method `join` to make two strings:

1. "Red Orange Yellow Green Blue Indigo and Violet" (call this one `without_commas`)
2. "Red, Orange, Yellow, Green, Blue, Indigo, and Violet" (call this one `with_commas`)

**Hint:** If you're still not sure what `join` does, first try just calling, for example, `"Hues".join(rainbow)` in a new code cell below.

```python
without_commas = ...
with_commas = ...

print('without commas: ', without_commas)
print('with commas: ', with_commas)
```

```python
grader.check("q1_3")
```

## 2. Indexing Arrays

These exercises give you practice accessing individual elements of arrays.  In Python (and in many programming languages), elements are accessed by *index*, so the first element is the element at index 0.

**Note:** If you have previous coding experience, you may be familiar with bracket notation. **DO NOT** use bracket notation when indexing (i.e. `arr[0]`), as this can yield different data type outputs than what we will be expecting. This can cause you to fail an autograder test.**

Be sure to refer to the [Python Reference](http://data8.org/sp22/python-reference.html) on the website if you feel stuck!

### Question 2.1

Create an array called `some_integers` that includes the intger 14, -3, -10, 27, 46, in that order. Set `second_element` to the second element of `some_integers` using the `item` method.

```python
some_integers = ...
second_element = ...
second_element
```

```python
grader.check("q2_1")
```

### Question 2.2

The next cell creates a table that displays some information about the elements of `some_integers` and their order.  Run the cell to see the partially-completed table, then fill in the missing information (the cells that say "Ellipsis") by assigning `blank_a`, `blank_b`, `blank_c`, and `blank_d` to the correct elements in the table.

**Hint:** Replace the ... with strings or numbers to make the Table consistent with the array from the previous problem.

```python
blank_a = ...
blank_b = ...
blank_c = ...
blank_d = ...

# Don't change anything below this line!
elements_of_some_integers = Table().with_columns(
    "English name for position", make_array(blank_a, "second", "third", blank_b, "fifth"),
    "Index",                     make_array(blank_c, 1, blank_d, 3, 4),
    "Element",                   some_integers)
elements_of_some_integers
```

```python
grader.check("q2_2")
```

### Question 2.3 

You'll sometimes want to find the *last* element of an array.  Suppose an array has 1024 elements.  What is the index of its last element?

```python
index_of_last_element = ...
```

```python
grader.check("q2_3")
```

More often than not, you don't know the number of elements in an array, also known as the array's *length*. For example, it might be a large data set you found on the Internet and aren't too familiar with. The function `len` takes a single argument, an array (or any collection of objects in Python), and returns the `len`gth of that array. The length will always be returned as an integer.

## 3. Basic Array Arithmetic

### Question 3.1

Divide the numbers `28`, `-2,882`, `2,828,242,428`, and `380` by `15`. Assign each variable below such that `first_quotient` is assigned to the result of $28 \div 15$, `second_quotient` is assigned to the result of $-2,882 \div 15$, and so on. 

For this question, **don't** use arrays.

```python
first_quotient = ...
second_quotient = ...
third_quotient = ...
fourth_quotient = ...
print(first_quotient, second_quotient, third_quotient, fourth_quotient)
```

```python
grader.check("q3_1")
```

### Question 3.2

Now, do the same calculation, but using an array called `numbers` that contains the 4 original numbers (`28`, `-2,882`, `2,828,242,428`, and `380`) and only a single division (`/`) operator.  Store the 4 results in an array named `quotients`.

```python
numbers = ...
quotients = ...
quotients
```

```python
grader.check("q3_2")
```

### Question 3.3

Oops, we made a typo!  Instead of dividing by 15, we wanted to **multiply** each number by 15.  Compute the correct products in the cell below using array arithmetic and store them in an array called `products`. The purpose of this question is to help you realize that it's much easier to perform (and in this case, change) arithmetic on an array all at once, rather than having to do so once for each element in the array. Notice that your job is going to be really easy if you previously defined an array containing the 4 numbers.

```python
products = ...
products
```

```python
grader.check("q3_3")
```

## 4. McDonald's Menu Items

![](coolmcdonalds.jpg)

McDonald's is recognized as one of the most recognizable and profitable fast-food restaurants in the world. With locations in over 100 countries, there is expected to be close to 44,000 McDonald's restaurants by the end of 2024. At it's inception, in 1940, McDonald's was a bar-b-q restaurant. It later reopened as a self-service drive-in with only nine menu items: hamburger, cheeseburger, soft drinks, milk, coffee, potato chips and a slice of pie. In the upcoming problems we will look at McDonald's menu items as well as their financial statements. 

The cell below loads a table called `menu_info` that includes a range of information about the nutritional value of McDonald's menu items.

```python
menu_info = Table.read_table("mcdonalds_data.csv")
menu_info
```

### Question 4.1

The cell below selects the column labeled `'calories'` from the `menu_info` Table and assigns it to the array called `item_calories`. This column contains the calories of McDonald's menu items. **Reminder:** Calling `.column(...)` on a table returns an array of the column specified, in this case the `calories` column of the `menu_info` table.

Use the `max` function to assign the largest calorie value found in the Table to `highest_calories`.

```python
item_calories = menu_info.column('calories')

highest_calories = ...
highest_calories
```

```python
grader.check("q4_1")
```

<!-- BEGIN QUESTION -->

### Question 4.2

There are several other ways you could have determined the maximum calorie value in the `menu_info` Table without using the `max` function. For example, there's one that uses both the `len(...)` and `.item(...)` functions, and you doesn't have to know the number of elements in `item_calories` to compute the value.

First, explain in words how you could use Table methods (for example, `.sort`), in combination with other array functions and methods to compute the maximum number of calories in the Table. Then, create a new code cell and implement your solution using Python code. Your code should be able to work no matter how big or small `item_calories` might be.

You can compare to this question to the answer in the previous question.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 4.3

It is recommended that some people eat 2,500 calories a day. We can see that one could easily exhaust a large percentage of their daily caloric intake on one McDonald's meal. Find the percent of daily caloric intake for each menu item based on a 2,500 calorie diet. Round to two decimal places.  Store the values in the array named, `percent_caloric_intake`.

```python
percent_caloric_intake = ...
percent_caloric_intake
```

```python
grader.check("q4_3")
```

### Question 4.4 

Select the column labeled `'totalfat'` from the `menu_info` Table and assign it to the array named `total_fat`. This column describes the total fat, measured in grams, of each McDonald's menu item. Each gram of fat contains 9 calories ([source](https://www.nal.usda.gov/programs/fnic)). 

Then, compute how many calories of each menu item come from macronutrients that are **not** fat (for example, sugar, protein, and carbohydrates), and assign those values to the array `nonfat_calories`.

```python
total_fat = ...

nonfat_calories = ...
nonfat_calories
```

```python
grader.check("q4_4")
```

## 5. McDonald's Financial Statements

The tests from this point on will **not** necessarily tell you whether or not your answers are correct.

The cell below loads a table of McDonald's financial statements for each year between 2002 and 2022. The data set comes from [Kaggle](https://www.kaggle.com/code/jayrdixit/mc-donalds-finance).

```python
mcdonalds_financials = Table.read_table("McDonalds_Financial_Statements.csv")
mcdonalds_financials
```

### `numpy` functions

In upcoming questions, you will apply some built-in `numpy` functions to this array. `numpy` is a Python module that is often used in Data Science. This homework will use two different `numpy` functions, `np.diff` and `np.cumsum`. **Note:** While used on this homework, these specific functions are not used heavily throughout the course. This set of questions is meant to get you quickly learning and using new functions in your programming, which is something data scientists need to do all the time!

#### `np.diff`
The difference function `np.diff` subtracts each element in an array from the element after it within the array. As a result, the length of the array `np.diff` returns will always be one less than the length of the input array.

The diagram below illustrates how this function operates.
<img src="array_diff.png" style="width: 600px;"/>

#### `np.cumsum`

The **cumulative** sum function `np.cumsum` outputs an array of partial sums. For example, the third element in the output array corresponds to the sum of the first, second, and third elements.

The diagram below illustrates how this function operates.

<img src="array_cumsum.png" style="width: 700px;"/>

### Question 5.1 

The first line below assigns `revenues` to an array of 21 revenues from consecutive years of McDonald's financial statements; specifically, 2002-2022. Assign the names `least`, `greatest`, and `average` so that the `print` statement is correct.

```python
years_revenue = Table.read_table("McDonalds_Financial_Statements.csv").sort("Year").column('Revenue ($B)')

least = ...
greatest = ...
average = ...

print("McDonald's revenue for 2002 to 2022 is between", least, "and", greatest, "billion dollars and", average, "billion dollars on average.")
```

```python
grader.check("q5_1")
```

### Question 5.2 

Assign `biggest_increase` to the biggest increase in revenue between two consecutive years. For example, the second year in the data set has a revenue of 17.14 billion dollars and the third of 19.06  billion dollars, so the increase in revenue was $19.06 - 17.14 = 1.92$ billion dollars, for an increase of 1.92 billion dollars.

**Hint:**

- You'll need an array arithmetic function [mentioned in the textbook](https://inferentialthinking.com/chapters/05/1/Arrays.html#functions-on-arrays). 


```python
biggest_increase = ...
biggest_increase
```

```python
grader.check("q5_2")
```

### Question 5.3

You determined in an earlier question that the average revenue during this time period was around 22.7 billion dollars. Therefore, you might *expect* the **total cumulative** revenue to be `22.7 * n` billion dollars after `n` years during this 21 year period.

Use the `np.arange` function to create an array of values where each element represents the **total cumulative** revenue (in billions of dollars) one would expect McDonald's to have made by the end of each year, for 21 years. Use the assumption that they make 22.7 billion dollars each year. For example, the first element would be 22.7, the second element 45.4, the third 68.1, and so on. Assign to this array of values to `expected_revenue`. You can read about `np.arange` as needed in this [textbook section](https://inferentialthinking.com/chapters/05/2/Ranges.html).

Then, determine how the actual total cumulative revenue over the 21 year period compared to what we expected to happen under our assumption. Set `difference_from_expected` to an array of 21 elements, where the element is the absolute difference between the actual cumulative revenue and expected cumulative revenue in that particular year. 

For example, since the first three revenues are 15.4, 17.14, and 19.06 billion dollars. The total cumulative revenue for 3 years is 15.4 + 17.14 + 19.06 = 51.6. The expected revenue for 3 years is 22.7 * 3 = 68.1. Therefore, `difference_from_expected.item(2)` should be $\lvert 51.6 - 68.1 \rvert = 16.5$. 

**Hint 1**: You'll need to use the `np.cumsum` function to create the array of actual cumulative revenues.

**Hint 2**: When using the absolute value function, use `np.abs` instead of `abs`, since `numpy` is already loaded in this notebook, and that's what the autograder is going to use.

```python
expected_revenue = ...
difference_from_expected = ...
difference_from_expected
```

```python
grader.check("q5_3")
```

### Question 5.4

Let’s imagine your guess for the next revenue was always just the revenue of the previous year. If you always guessed the previous year's revenue, how big would your error in guessing the revenue be, on average? Compute the value using some `numpy` operations and assign your answer to `average_error`.

For example, the first three revenues are 15.4, 17.14, and 19.06. After the first year you guess the next year's revenue will be 15.4 billion dollars, but it is 17.14 billion dollars resulting in an error of $\lvert 15.4- 17.14 \rvert$. For the next year you now guess there will be a revenue of 17.14 billion dollars, but there is a revenue of 19.06 billion dollars, so the error is $\lvert 17.14 - 19.06 \rvert$. Therefore, the average difference between your guess and the actual revenue for just the second and third years would be $$\frac{|15.4-17.14|+ |17.14-19.06|}{2} = 1.83$$.

```python
average_error = ...
average_error
```

```python
grader.check("q5_4")
```

## 6. Tables

### Question 6.1 
![](fruitstand.png)

Suppose you have 4 apples, 3 oranges, and 3 pineapples.  (Perhaps you're using Python to solve a high school Algebra problem.)  Create a table that contains this information.  It should have two columns: `fruit name` and `count`.  Assign the new table to the variable `fruits`.

**Note:** Use lower-case and singular words for the name of each fruit, like `"apple"`.

Your table should look like:

| fruit name | count |
|------------|-------|
| apple      | 4     |
| orange     | 3     |
| pineapple  | 3     |


```python
# Our solution uses 1 statement split over 3 lines. 
# You can write yours on 1 line if you wish.
fruits = ...
    ...
    ...
fruits
```

```python
grader.check("q6_1")
```

### Question 6.2

The file `inventory.csv` contains information about the inventory at a fruit stand.  Each row represents the contents of one box of fruit. Load it as a table named `inventory` using the `Table.read_table()` function. `Table.read_table(...)` takes one argument (data file name in string format) and returns a table.

```python
inventory = ...
inventory
```

```python
grader.check("q6_2")
```

### Question 6.3 

The file `sales.csv` contains the number of fruit sold from each box last Saturday.  It has an extra column called "price per fruit (\$)" that's the price *per item of fruit* for fruit in that box.  The rows are in the same order as the `inventory` table.  

Load these data into a table called `sales`.

```python
sales = ...
sales
```

```python
grader.check("q6_3")
```

### Question 6.4 

How many fruits did the store sell in total on that day?

```python
total_fruits_sold = ...
total_fruits_sold
```

```python
grader.check("q6_4")
```

### Question 6.5

What was the store's total revenue (the total price of all fruits sold) on that day?

**Hint:** If you're stuck, think first about how you would compute the total revenue from just the grape sales.

```python
total_revenue = ...
total_revenue
```

```python
grader.check("q6_5")
```

### Question 6.6

Make a new table called `remaining_inventory`.  It should have the same rows and columns as `inventory`, except that the amount of fruit sold from each box should be subtracted from that box's **original** count, so that the `count` column is **updated to be** the amount of fruit remaining after Saturday. 

The provided solution spans multiple lines, but you can write yours on a single line if you prefer.

```python
remaining_inventory = ...
    ...
    ...
    ...
remaining_inventory
```

```python
grader.check("q6_6")
```

# Submitting your work
You're done with this assignment! Assignments should be turned in using the following best practices:
1. Save your notebook.
2. Restart the kernel and run all cells up to this one.
3. Run the cell below with the code `grader.export(...)`. This will re-run all the tests. Make sure they are passing as you expect them to.
4. Download the file named `hw02_<date-time-stamp>.zip`, found in the explorer pane on the left side of the screen. **Note**: Clicking on the link in this notebook may result in an error, it's best to download from the file explorer panel.
5. Upload `hw02_<date-time-stamp>.zip` to the corresponding assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1_1 (points: None)
```
>>> isinstance(rainbow, np.ndarray)
True
```
failure_message: ❌ You did not create an array. Perhaps you created a list or another type of collection?
success_message: ✅ You successfully created an array.
```
>>> not any([',' in text for text in rainbow])
True
```
failure_message: ❌ You included commas inside some of the strings. Commas are used to seperate individual strings, but are not between the quotes.
```
>>> 'and ' in rainbow.item(6)
True
```
failure_message: ❌ Your last string does not contain the substring `and `
success_message: ✅ The last string contains the substring `and `
```
>>> len(rainbow) == 7
True
```
failure_message: ❌ The array named rainbow does not have 7 elements
success_message: ✅ The array named rainbow has 7 elements


Tests for q1_2 (points: None)
```
>>> isinstance(just_numbers, np.ndarray)
True
```
failure_message: ❌ You did not create an array. Perhaps you created a list or another type of collection?
success_message: ✅ You successfully created an array.
```
>>> len(just_numbers)
4
```
failure_message: ❌ Your array does not have the correct number of elements in it.
success_message: ✅ Your array has the correct number of elements in it.


Tests for q1_3 (points: None)
```
>>> ',' in with_commas
True
```
failure_message: ❌ The string `with_commas` does not have any commas in it
success_message: ✅ The string `with_commas` has at least one comma in it
```
>>> ',' not in without_commas
True
```
failure_message: ❌ The string `without_commas` has at least one comma in it
success_message: ✅ The string `without_commas` does not have any commas in it
```
>>> len(with_commas)
52
```
failure_message: ❌ The string `with_commas` is not the correct length
success_message: ✅ The string `with_commas` is the correct length
```
>>> len(without_commas)
46
```
failure_message: ❌ The string `without_commas` is not the correct length
success_message: ✅ The string `without_commas` is the correct length


Tests for q2_1 (points: None)
```
>>> np.array_equal(some_integers, np.array([14, -3, -10, 27, 46]))
True
```
failure_message: ❌ You did not sucessfully create the `some_integers` array
success_message: ✅ You sucessfully created the `some_integers` array
```
>>> second_element in some_integers
True
```
failure_message: ❌ The number you assigned to `second_element` is not from the array `some_numbers`
success_message: ✅ The number you assigned to `second_element` is from the array `some_numbers`
```
>>> second_element != -10
True
```
failure_message: ❌ You incorrectly selected the element at index 2. This is the fourth element of the array.


Tests for q2_2 (points: None)
```
>>> elements_of_some_integers.column(0).item(2)
'third'
```
failure_message: ❌ `blank_a` is incorrect. Check your answer and the spelling. There should be no Punctuation or spaces.
success_message: ✅ `blank_a` is correct
```
>>> elements_of_some_integers.column(0).item(3)
'fourth'
```
failure_message: ❌ `blank_b` is incorrect. Check your answer and the spelling. There should be no Punctuation or spaces.
success_message: ✅ `blank_b` is correct
```
>>> elements_of_some_integers.column(1).item(0)
0
```
failure_message: ❌ `blank_c` is incorrect. Check your answer and the spelling. There should be no Punctuation or spaces.
success_message: ✅ `blank_c` is correct
```
>>> elements_of_some_integers.column(1).item(3)
3
```
failure_message: ❌ `blank_d` is incorrect. Check your answer and the spelling. There should be no Punctuation or spaces.
success_message: ✅ `blank_d` is correct


Tests for q2_3 (points: None)
```
>>> index_of_last_element
1023
```
failure_message: ❌ You incorrectly determined the index. Don't forget, indexing starts at 0.
success_message: ✅ You correctly determined the index


Tests for q3_1 (points: None)
```
>>> first_quotient
1.8666666666666667
```
failure_message: ❌ You incorrectly computed `first_quotient`. Double check you've typed the numbers correctly!
success_message: ✅ You correctly computed `first_quotient`.
```
>>> second_quotient
-192.13333333333333
```
failure_message: ❌ You incorrectly computed `second_quotient`. Double check you've typed the numbers correctly!
success_message: ✅ You correctly computed `second_quotient`.
```
>>> third_quotient
188549495.2
```
failure_message: ❌ You incorrectly computed `third_quotient`. Double check you've typed the numbers correctly!
success_message: ✅ You correctly computed `third_quotient`.
```
>>> fourth_quotient
25.333333333333332
```
failure_message: ❌ You incorrectly computed `fourth_quotient`. Double check you've typed the numbers correctly!
success_message: ✅ You correctly computed `fourth_quotient`.


Tests for q3_2 (points: None)
```
>>> np.array_equal(numbers, np.array([28, -2882, 2828242428, 380]))
True
```
failure_message: ❌ You did not create the `numbers` array. Perhaps you created a list or another type of collection?
success_message: ✅ You successfully created the `numbers` array.
```
>>> type(quotients) == np.ndarray
True
```
failure_message: ❌ You did not create an array named `quotients`. Perhaps you created a list or another type of collection?
success_message: ✅ You successfully created an array named `quotients`.


Tests for q3_3 (points: None)
```
>>> np.allclose(products, np.array([420, -43230, 42423636420, 5700]))
True
```
failure_message: ❌ Your array does not contain the correct elements, or they are in an incorrect position.
success_message: ✅ Your array has the correct elements in it.


Tests for q4_1 (points: None)
```
>>> highest_calories
834.36000000000001
```
failure_message: ❌ You did not find the menu item with the highest number of calories
success_message: ✅ You correctly found the menu item with the highest number of calories


Tests for q4_3 (points: None)
```
>>> len(percent_caloric_intake)
141
```
failure_message: ❌ Your array has an incorrect number of etnries
success_message: ✅ Your array has the correct number of entries
```
>>> percent_caloric_intake.item(25)
16.44
```
failure_message: ❌ At least one element of your array is incorrect
success_message: ✅ At least one element of your array is correct
```
>>> np.round(sum(percent_caloric_intake)) == 1367.0 and np.round(np.std(percent_caloric_intake)) == 7.0
True
```
failure_message: ❌ Your array doesn't appear to contain the correct elements
success_message: ✅ Your array contains the correct elements


Tests for q4_4 (points: None)
```
>>> len(nonfat_calories)
141
```
failure_message: ❌ Your array has an incorrect number of etnries
success_message: ✅ Your array has the correct number of entries
```
>>> nonfat_calories.item(0)
277.53
```
failure_message: ❌ The first element of your array is incorrect
success_message: ✅ The first element of your array is correct
```
>>> np.round(sum(nonfat_calories), 0) == 21411.0 and np.round(np.std(nonfat_calories), 0) == 108.0
True
```
failure_message: ❌ Your array doesn't appear to contain the correct elements
success_message: ✅ Your array contains the correct elements


Tests for q5_1 (points: None)
```
>>> least <= average <= greatest
True
```
failure_message: ❌ The values you computed aren't ordered correctly
success_message: ✅ least <= average <= greatest


Tests for q5_2 (points: None)
```
>>> 3 <= biggest_increase < 5
True
```
failure_message: ❌ `biggest_increase` is not close to the correct answer
success_message: ✅ `biggest_increase` is at least in the ballpark


Tests for q5_3 (points: None)
```
>>> len(difference_from_expected)
21
```
failure_message: ❌ `difference_from_expected` doesn't have the correct number of elements
success_message: ✅ `difference_from_expected` has the correct number of elements


Tests for q5_4 (points: None)
```
>>> 1 <= average_error <= 3
True
```
failure_message: ❌ `average_error` is not close to the correct answer
success_message: ✅ `average_error` is at least in the ballpark


Tests for q6_1 (points: None)
```
>>> type(fruits) == Table
True
```
failure_message: ❌ You did not create a Table object
success_message: ✅ You created a Table and assigned it to `fruits`
```
>>> fruits.labels == ('fruit name', 'count')
True
```
failure_message: ❌ Your Table has incorrect column labels or they are in the incorrect order
success_message: ✅ Your Table has the correct column labels in the correct order


Tests for q6_2 (points: None)
```
>>> inventory.sort(0)
box ID | fruit name | count
25274  | apple      | 20
26187  | strawberry | 255
43566  | peach      | 40
48800  | orange     | 35
52357  | strawberry | 102
53686  | kiwi       | 45
57181  | strawberry | 123
57930  | grape      | 517
```
failure_message: ❌ You did not correctly load `inventory.csv`
success_message: ✅ You loaded `inventory.csv` correctly


Tests for q6_3 (points: None)
```
>>> sales.sort(0)
box ID | fruit name | count sold | price per fruit ($)
25274  | apple      | 0          | 0.8
26187  | strawberry | 25         | 0.15
43566  | peach      | 17         | 0.8
48800  | orange     | 35         | 0.6
52357  | strawberry | 102        | 0.25
53686  | kiwi       | 3          | 0.5
57181  | strawberry | 101        | 0.2
57930  | grape      | 355        | 0.06
```
failure_message: ❌ You did not correctly load `sales.csv`
success_message: ✅ You loaded `sales.csv` correctly


Tests for q6_4 (points: None)
```
>>> type(total_fruits_sold) == np.int64 or type(total_fruits_sold) == int
True
```
failure_message: ❌ Your answer should be an integer
success_message: ✅ You provided an integer answer
```
>>> total_fruits_sold > 500
True
```
failure_message: ❌ Your number is not large enough to be the total number of pieces of fruits sold
success_message: ✅ Your number is large enough to be reasonable as an answer


Tests for q6_5 (points: None)
```
>>> 50 <= total_revenue <= 150
True
```
failure_message: ❌ If you're stuck, here's a hint. You want to multiply the count sold in each box by the per-item price of fruits in that box, and then sum up each product.
success_message: ✅ Your number is between 50 and 150 dollars, which means it is a reasonable answer


Tests for q6_6 (points: None)
```
>>> remaining_inventory.num_columns
3
```
failure_message: ❌ Your table doesn't have all 3 columns
success_message: ✅ Your table has the correct number of columns
```
>>> remaining_inventory.column('count').item(0) != 45
True
```
failure_message: ❌ You forgot to subtract off the sales
success_message: ✅ You remembered to subtract off the sales
```
>>> np.array_equal(remaining_inventory.sort(0).column('fruit name'), ['apple', 'strawberry', 'peach', 'orange', 'strawberry', 'kiwi', 'strawberry', 'grape'])
True
```
failure_message: ❌ You forgot to subtract off the sales
